# Main task
Let's start with loading the library we need.

```python
import re
```

### Part 1: names

<!-- grade_id: cell-9093ba874d8b728b, points: 0, grade: false, solution: true -->
```python
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""

    
    # YOUR CODE HERE

names()
```

<!-- grade_id: cell-b4abb40ddaf85113, points: 1, grade: true, solution: false -->
```python
### Please do not type anything in this cell.
assert len(names()) == 4, "There are four names in the simple_string"
assert ((len(names()[0])>2)&(len(names()[0])>2)&(len(names()[0])>2)&(len(names()[0])>2)) == True, "Please match all the characters in these names."
assert ('Amy' in names())== True
assert ('Mary' in names())== True
assert ('Ruth' in names())== True
assert ('Peter' in names())== True

```

### Part 2: student grades
The dataset file in assets/grades.txt contains multiple lines of people along with their grades in a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: cell-1f815af992a50638, points: 0, grade: false, solution: true -->
```python
def student_grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()
    
    # YOUR CODE HERE

student_grades()
```

<!-- grade_id: cell-06db205fa3b40c65, points: 1, grade: true, solution: false -->
```python
### Please do not type anything in this cell.
assert len(student_grades()) == 16
name_list = ['Bell Kassulke',
 'Simon Loidl',
 'Elias Jovanovic',
 'Hakim Botros',
 'Emilie Lorentsen',
 'Jake Wood',
 'Fatemeh Akhtar',
 'Kim Weston',
 'Yasmin Dar',
 'Viswamitra Upandhye',
 'Killian Kaufman',
 'Elwood Page',
 'Elodie Booker',
 'Adnan Chen',
 'Hank Spinka',
 'Hannah Bayer']

assert (': B' in student_grades()[1]) == False, "Please only return student names."
assert ('Kassulke' in student_grades()) == False, "Please return students' full name."
for name in name_list:
    assert (name in student_grades())== True, "Check your name list."

```

### Part 3: logs
Consider the variable 'logdata' which is a string containing a standard web log. This variable records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items: 

* a host (e.g., ‘146.204.224.152’) 
* a user (e.g., ‘feest6811’) 
* the time a request was made (e.g., ‘21/Jun/2019:15:45:24 -0700’) 
* the post request type (e.g., ‘POST /incentivize HTTP/1.1’)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:

```
example_dict = {"host":"146.204.224.152",
                "user":"feest6811",
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```



<!-- grade_id: cell-b0fa9dbee5089ac9, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    
    # YOUR CODE HERE

logs()
```

<!-- grade_id: cell-ca9d4edf008ce4ff, points: 1, grade: true, solution: false -->
```python
### Please do not type anything in this cell.
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

assert len(logs()[0:1][0])==4, "Make sure that each of your dictionary is in the right format. "

tmp3 = [{'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'},
 {'host': '98.143.107.61',
  'user_name': 'reilly5813',
  'time': '21/Jun/2019:15:54:04 -0700',
  'request': 'HEAD /channels HTTP/1.1'},
 {'host': '93.88.34.132',
  'user_name': '-',
  'time': '21/Jun/2019:15:49:04 -0700',
  'request': 'POST /24%2f365 HTTP/2.0'},
 {'host': '100.32.205.59',
  'user_name': 'ortiz8891',
  'time': '21/Jun/2019:15:45:28 -0700',
  'request': 'PATCH /architectures HTTP/1.0'},
 {'host': '168.95.156.240',
  'user_name': 'stark2413',
  'time': '21/Jun/2019:15:45:31 -0700',
  'request': 'GET /engage HTTP/2.0'},
 {'host': '72.157.89.10',
  'user_name': '-',
  'time': '21/Jun/2019:15:47:14 -0700',
  'request': 'PATCH /clicks-and-mortar/metrics HTTP/1.0'},
 {'host': '180.95.121.94',
  'user_name': 'mohr6893',
  'time': '21/Jun/2019:15:45:34 -0700',
  'request': 'PATCH /extensible/reinvent HTTP/1.1'},
 {'host': '144.23.247.108',
  'user_name': 'auer7552',
  'time': '21/Jun/2019:15:45:35 -0700',
  'request': 'POST /extensible/infrastructures/one-to-one/enterprise HTTP/1.1'},
 {'host': '180.234.185.134',
  'user_name': 'windler5463',
  'time': '21/Jun/2019:15:46:34 -0700',
  'request': 'POST /synergies/innovative/enterprise HTTP/1.1'},
 {'host': '233.187.15.207',
  'user_name': '-',
  'time': '21/Jun/2019:15:46:14 -0700',
  'request': 'GET /harness/intuitive HTTP/1.0'}]

for dic in tmp3:
    assert dic in logs(), "Please include all the records"

```